# CS 101:  Introduction to Computing for Science and Engineering

## `lab03`—Types, Functions, & Scope

**Objectives**

-   Use index slicing to access larger pieces of a string.
-   Use functions to modularize code.
-   Utilize the built-in libraries of Python functions effectively.

#### Slicing

We can generalize our notion of *indexing* to *slicing*.  Slicing means using a *range of indices* to take a *slice* of a string.  A range is written `string[left:right]`, where the `left` index is part of the string but the `right` index is not.  (We call this a *half-open* range.)

```python
pangram = 'Sphinx of black quartz: judge my vow.'
```

```python
pangram[0:5]
```

Output:
```
'Sphin'
```

```python
pangram[6:13]
```

Output:
```
' of bla'
```

```python
pangram[0:0]  # Note that the right-hand index is not in the range, so this string is empty.
```

Output:
```
''
```

```python
pangram[0:-1]
```

Output:
```
'Sphinx of black quartz: judge my vow'
```

```python
pangram[5:]  # We don't have to specify both sides; if we don't, Python assumes we are going all the way to the end or coming all the way from the beginning.
```

Output:
```
'x of black quartz: judge my vow.'
```

```python
pangram[:10]
```

Output:
```
'Sphinx of '
```

```python
pangram[:]  # Thus this is just a copy of the entire original string.
```

Output:
```
'Sphinx of black quartz: judge my vow.'
```

```python
pangram[4:-1:3]  # Take every THIRD character from the fourth space to the last (exclusive).
```

Output:
```
'nobcqr:ueyo'
```

Although strings *are* important and useful, you're still fairly limited in what you're able to do with them.  But many aspects of indexing will carry over to data types you'll meet later, like `list`s and `ndarray`s ($n$-dimensional arrays)—you will use these techniques to manipulate large arrays of numerical data as well.

-   How would you extract the letters `'ssupdla nioiaim'` from the following string?  Write the full expression (*i.e.*, `showoff[???]`) in the answer block below.
        
        showoff = 'sesquipedalian antinominalism'

```python
# definition
showoff = 'sesquipedalian antinominalism'
```

<!-- grade_id: slice, points: 1, grade: true, solution: true -->
```python
# your answer here
showoff[0] + showoff[0] + showoff[4] + showoff[3] + showoff[8] + showoff[10] + showoff[9] + showoff[14] + showoff[-6] + showoff[-3] + showoff[-9] + showoff[-3] + showoff[-5] + showoff[-3] + showoff[-1]
```

Output:
```
'ssuqdla nioiaim'
```

### How to reuse your code

When you need to carry out a calculation many times (whether simple or complicated), you don't want to write out the formula over and over again.  Besides being inefficient, if you later need to expand the code or correct a bug, you'd have to look everywhere to make sure you changed every case.  It's far better to write the calculation *once*, as a function, and then call that function every time you need it.

You've seen a handful of functions before:  `type`, `exp`, `abs`.  Let's see what's involved with writing your own.

In mathematics, a function is simply a way of relating an input value (often $x$) to an output value (sometimes $y$ or $f(x)$).  For instance, a parabola may be expressed as

$$y = f(x) = x^2$$

and plotted as

![](./img/parabola.png)

Properly speaking, the resulting curve represents the collection of ordered pairs $(x,y)$.  To plot this, we select a discrete set of points in the domain $x$ and calculate the corresponding set of values $y$.

| $x$  | -1.5 | -1.25  | -1    | -0.75 | -0.5  | -0.25 | 0  | .25    | 0.5  | 0.75   | 1 | 1.25  |  1.5  |
|--------|------|--------|-------|-------|-------|-------|----|--------|------|--------|---|-------|-------|
| $y$ | 2.25 | 1.5625 | 1.0 | 0.5625 | 0.25 | 0.0625 |  0 | 0.0625 | 0.25 | 0.5625 | 1 | 1.5625 | 2.25 |

Of course, the advantage of using a computer is that each step is easier to type than to calculate by hand.  If you had to fill in this table, even using Python it would be tedious if you had to reset the variable `x` and retype the expression for `y` by hand every time.

It is very likely that any equation you will wish to use in your studies and career will be more complicated than $y = x^{2}$.  In order to minimize the amount of time typing code, as well as the likelihood of making a mistake due to repetition, we can use functions to capture the operations involved:

```python
def sqr(x):
    result = x * x
    return result
```

Now instead of typing `x**2` or `x*x` each time you need the formula, you can call the function `sqr(x)` instead.  (In this trivial case, that's actually *more* keystrokes, but the principle stands.)

A more realistic candidate is the quadratic equation (as we saw implemented in lab2), or the equation for an arbitrary line, $y = mx + b$, where $m$ is the slope and $b$ is the intercept.  A function which formats a slope and intercept into a string representing a line may look like this:

```python
def line_eqn(m,b):
    return 'y = %f x + %f'%(m,b)
```

```python
print(line_eqn(1.875,4))
```

Output:
```
y = 1.875000 x + 4.000000
```

#### `math` Library Functions

*Modules* are collections of functions which allow you to take advantage of the collective work of the scientific and engineering community when carrying out your calculations or building your models.  One of the strengths of Python is its enormous and rich library of modules, which provide both basic and involved operations far beyond the built-in operators.

Let's consider the `math` library.  To use functions like `sin`, `cos`, `exp`, and `log`, you simply need to `import math`:

```python
import math
```

You still need to specify *which* module contains the functions to use them, but otherwise it's quite straightforward:

$$
\sin\left(\frac{\pi}{4}\right)
$$

```python
math.sin(math.pi/4)
```

Output:
```
0.7071067811865475
```

$$
\left[\text{ceil}\left(4.5\right)\right] !
$$

```python
math.factorial(math.ceil(4.5))
```

Output:
```
120
```

-   Write an expression for the following formula:
    
    $$
    \frac{(m+n)!}{(m+1)(n-1)!}
    $$
    
    (Note that the factorial refers only to the second term of the denominator, not the entire denominator.)

```python
# these are representative values—don't change these
m = 2
n = 3
```

<!-- grade_id: form, points: 0, grade: false, solution: true -->
```python
# write your expression here
expr = math.factorial((m + n))/((m + 1)*(math.factorial((n-1))))# your code here
```

<!-- grade_id: form-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(expr, 20.0)
print('Success!')
```

Output:
```
Success!
```

Some modules are quite large, and you may not want to import everything into memory at once.  In that case, you can import a single function from a module as follows:

```python
from math import asin  # arcsine
```

-   What line of Python code could you write to import `log10` from `math`?  (This is the logarithm base 10, $\log_{10}(x)$.)

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
from math import log10# write your expression here
```

<!-- grade_id: import-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
try:
    log10(0.25)
except NameError:
    raise AssertionError('log10 not imported yet')
print('Success!')
```

Output:
```
Success!
```

### Application:  Cardinal Sine Function

The bread and butter of working effectively in Python is writing your own functions.

Consider the *cardinal sine function*, $\text{sinc}(x)$, which is used in signal processing by electrical engineers.

$$\text{sinc}(x) = \frac{\sin(x)}{x}$$

![](https://upload.wikimedia.org/wikipedia/commons/thumb/5/59/Si_sinc.svg/700px-Si_sinc.svg.png)

To compose a function which reproduces this behavior for an arbitrary number, we could simply define:

```python
from math import sin
def sinc(x):
    y = sin(x) / x
    return y
```

This short function simply creates a new variable `y` which contains the result of `sin(x)/x` evaluated for the actual value of `x`.  The `return` statement tells Python what the function gives back—in this case, the value of `y`.

```python
sinc(1.5)
```

Output:
```
0.6649966577360363
```

```python
sinc(3.1415)
```

Output:
```
2.949342341572187e-05
```

Critically, *the name of the variable OUTSIDE of `sinc` is not related to the name of the variable INSIDE `sinc`*.  (We say that they inhabit different *scopes*.)  Consider this:

```python
y = 5.0
sinc(y)
```

Output:
```
-0.1917848549326277
```

-   For what value of input `x` does `sinc` fail?  (You can find this by experiment or inspection.)

<!-- grade_id: divbyzero, points: 1, grade: true, solution: true -->
```python
(0.0)# your answer here, in the form of a float (e.g., 1.0)
```

Output:
```
0.0
```

<div class="alert alert-warning">
<p>
The `y` in the scope outside of the function is equal to `5.0`, whereas inside of the function `sinc` the same *value* is referred to as `x`, and `y` means something else.
</p><p>
You can call variables any valid name you like, but it's generally good practice to name it something clear in context.  That is,
    
    <pre><code>def sinc(number_of_calabazas):
        return sin(number_of_calabazas) / number_of_calabazas
</code></pre>
is a little too bohemian to represent good programming style.  **Responsible variable names reflect their use.**
</p>
</div>

### Application:  Atmospheric Pressure

A mercury barometer is frequently used to precisely measure the atmospheric pressure.  The height $h$ of a column of mercury adjusts to a point where its weight plus force due to the mercury vapor pressure $p_{\text{vapor}}$ balances the force due to the external atmospheric pressure $p_{\text{atm}}$.  Mathematically, we write,

$$p_{\text{atm}} = \gamma h + p_{\text{vapor}}$$

where $\gamma$ is the specific weight of mercury.  In SI units, $\gamma = 133 \times 10^{3} \,\text{N}\cdot\text{m}^{-3}$.  The vapor pressure of mercury is typically very small in relation to the other values, so we can safely neglect it ($\left. p_{\text{vapor}} \right|_{T=25 ^{\circ}\text{C}} = 26.13 Pa$)<sup>[[NISTIR 6643](http://www.boulder.nist.gov/div838/SelectedPubs/NISTIR.6643.pdf)]</sup>.  A function which calculates the atmospheric pressure given the height of a column of mercury, then, is:

```python
def p_pa(h_Hg):
    # Calculate and return the atmospheric pressure (in Pa) given the height
    # of a column of mercury (in m).
    gamma = 133e3       # N/m^-3
    P = gamma * h_Hg    # Pa
    return P
```

```python
p_pa(0.598)  # The atmospheric pressure in pascals for 598 mmHg.
```

Output:
```
79534.0
```

```python
p_pa(0.598) / 101325  # The atmospheric pressure in the units of standard atmosphere for 598 mmHg.
```

Output:
```
0.7849395509499136
```

We will now introduce the `input` function, which lets you supply a message and ask the user for a value.  (This value then needs to be converted from a string to a `float`.)

```python
h_str = input('Give the height of the column of mercury in m:')
h = float(h_str) # convert to a floating-point number
print('The result is %f pascals.' % p_pa(h))
```

Output:
```
Give the height of the column of mercury in m:.5
The result is 66500.000000 pascals.
```

-   Create a function `p_atm` which accepts the height of a column of mercury and returns the atmospheric pressure in standard atmospheres (demonstrated a couple of cells above this).  (This is equal to the pressure in pascals divided by $101\,325\,\text{Pa}/\text{atm}$.  You may accomplish this by either calling `p_pa` and dividing by 101325 or by doing the calculation directly.)

<!-- grade_id: mercury, points: 0, grade: false, solution: true -->
```python
def p_atm(h):
    gamma= 133e3
    P = (gamma* h)/101325
    return P# define your function here---you can copy and paste from above, or write your own from scratch
```

<!-- grade_id: mercury-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(p_atm(1.00), 1.3126079447322971)
print('Success!')
```

Output:
```
Success!
```

### Application:  The U.S.D.O.T. Bridge Formula

The United States Department of Transportation, as part of its stewardship of the U.S. interstate highway system, maintains code standards for bridge construction.  As a metric for preserving shipping infrastructure, the DOT uses the bridge formula<sup>[[DOT2014](http://ops.fhwa.dot.gov/freight/sw/brdgcalc/calc_page.htm)]</sup> to limit the legal weight-to-length ratio for freight-laden vehicles crossing a bridge.  In simplified form, the bridge formula is

$$W = 500 \left( \frac{LN}{N-1} + 12N + 36 \right)$$

where

-   $W$ is the overall gross weight on any group of consecutive axles;
-   $L$ is the distance in feet between the outer axles of any group of consecutive axles; and
-   $N$ is the number of axles in the group under consideration.

This prevents short trucks with heavy loads from damaging bridges.

![](http://ops.fhwa.dot.gov/Freight/publications/brdg_frm_wghts/images/fig1.gif)

You will implement this formula as a function and use it to answer the following questions.

-   Complete the function `bridge_formula` by adding lines of code to accomplish both tasks indicated.

<!-- grade_id: dot, points: 0, grade: false, solution: true -->
```python
def bridge_formula(L, N):
    '''
    Calculate bridge weight in accordance with federal law.

    Args:
        L:  distance between groups of consecutive axles, ft
        N:  number of axles in group under consideration
    Returns:
        W:  overall gross weight on any group of consecutive axles, lb
    '''

    # Calculate the bridge formula's weight W using L and N.
    W = 500*(((L*N)/(N-1)) + 12*N + 36)# ... your code here ...

    # Return the resulting weight W.
    return W# ... your code here ...
```

As a test of your function, the following values are correct:  $L = 51$, $N = 5$, $W = 79,500$.  (This describes a $51'$, 5-axle semi which is legally allowed to carry $79,500 \,\text{lb}$, including freight and truck weight—in practice, U.S.D.O.T. rounds to $80,000$ but we won't worry about that layer of complexity here.)

<!-- grade_id: dot-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert bridge_formula(51, 5) == 79875
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: dot-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert bridge_formula(20,3) == 51000
print('Success!')
```

Output:
```
Success!
```

### How to divide your code logically

```python
def warn(msg):
    print('WARNING:  ', msg)

warn('Keyboard not responding. Press any key to continue.')
```

Output:
```
WARNING:   Keyboard not responding. Press any key to continue.
```

### How to manipulate strings

One of the elegant notions of programming is that programming helps generalize the idea of function beyond just manipulating numbers.  What about other data types—in particular, the other data type we've examined thus far, the string?  Python does have a lot of basic string manipulation functions.

```python
myStr = 'Mr. Jock, TV quiz PhD, bags few lynx.  '
```

```python
len(myStr)              # counts number of characters in myStr
```

Output:
```
39
```

```python
myStr.count('x')        # counts number of occurrences of 'x' in myStr
```

Output:
```
1
```

```python
myStr.find('x')         # returns position of character 'x'
```

Output:
```
35
```

```python
myStr.lower()           # returns myStr as lower-case letters
```

Output:
```
'mr. jock, tv quiz phd, bags few lynx.  '
```

```python
myStr.upper()           # returns myStr as upper-case letters
```

Output:
```
'MR. JOCK, TV QUIZ PHD, BAGS FEW LYNX.  '
```

```python
myStr.title()           # returns myStr as title-case letters
```

Output:
```
'Mr. Jock, Tv Quiz Phd, Bags Few Lynx.  '
```

```python
myStr.replace('a', 'b') # replaces all occurrences of a with b in myStr
```

Output:
```
'Mr. Jock, TV quiz PhD, bbgs few lynx.  '
```

```python
myStr.strip()           # removes leading/trailing white space from myStr
```

Output:
```
'Mr. Jock, TV quiz PhD, bags few lynx.'
```

Note that `myStr` itself never changes—functions like `upper` *return a copy* of `myStr` with the requested change made.

String methods are also *composable*—that is, you can put them together to carry out more complex operations:

```python
myStr.strip().replace(' ', '-')
```

Output:
```
'Mr.-Jock,-TV-quiz-PhD,-bags-few-lynx.'
```

```python
myStr.replace(' ', '-').strip()
```

Output:
```
'Mr.-Jock,-TV-quiz-PhD,-bags-few-lynx.--'
```

Now answer these questions *using the string methods given above*:

```python
quip = '"Sometimes I wonder whether the world is being run by smart people who are putting us on, or by imbeciles who really mean it." (Mark Twain)'
```

-   How long is `quip`?

<!-- grade_id: quip1, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = len(quip)
```

-   How many times does `'b'` occur in this phrase?

<!-- grade_id: quip2, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = quip.count('b')
```

-   How many times does `' '` occur?

<!-- grade_id: quip3, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = quip.count(' ')
```

-   Replace all instances of `' '` with `'a'`.  Now how many times does `'a'` occur?  (You can do this using an intermediate variable or by composing the operations as above.)

<!-- grade_id: quip4, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = quip.replace(' ','a').count('a')
```

Finally, let's close out this section with a teaser for the `list`, which we discuss in lecture 7.  Try these:

```python
quip.split(' ')
```

Output:
```
['"Sometimes',
 'I',
 'wonder',
 'whether',
 'the',
 'world',
 'is',
 'being',
 'run',
 'by',
 'smart',
 'people',
 'who',
 'are',
 'putting',
 'us',
 'on,',
 'or',
 'by',
 'imbeciles',
 'who',
 'really',
 'mean',
 'it."',
 '(Mark',
 'Twain)']
```

```python
myStr.split(',')
```

Output:
```
['Mr. Jock', ' TV quiz PhD', ' bags few lynx.  ']
```

These break the string along the fault lines specified, or the *delimiters* `' '` and `','`, and return the result as a collection of strings known as a list.

-   Convert the string `' Time present and time past / Are both perhaps present in time future '` into the string `'Tim3_Pr3S3Nt_And_Tim3_Past_/_Ar3_Both_P3Rhaps_Pr3S3Nt_In_Tim3_Futur3'` using only the above functions.  You should only require four functions, but may take more if necessary.  The final result should be stored in the variable `ans`.

<!-- grade_id: xform, points: 0, grade: false, solution: true -->
```python
# your code here
start = ' Time present and time past / Are both perhaps present in time future '
ans = start.strip().replace('e','3').replace(' ','_').title()
```

<!-- grade_id: xform-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert ans == 'Tim3_Pr3S3Nt_And_Tim3_Past_/_Ar3_Both_P3Rhaps_Pr3S3Nt_In_Tim3_Futur3'
print('Success!')
```

Output:
```
Success!
```


# Before you submit...

Before you submit this exercise, make sure that everything runs as expected.  You can either use the *Validate* button on the *Assignments* tab or type `nbgrader validate labX` at the command line.  **You should also *save* before submitting.**

Make sure that you have filled in any code block which says "YOUR CODE HERE" or "YOUR ANSWER HERE".

List any resources used (beyond the TA or the help files):

This laboratory exercise was written by Neal Davis for Computer Science 101 at the University of Illinois.

©2015–16 University of Illinois

![](./cs101-footer.png)