# CS 101:  Introduction to Computing for Science and Engineering

## `lab06`—Data Analysis

**Objectives**

-   Use lists and list features to process data from a file.
-   Apply the standard pipeline of data analysis:  data cleaning and preparation, data processing, and output.

![](https://d1b10bmlvqabco.cloudfront.net/attach/ipzxix9y8ou155/idiy91xwl663a4/ituwx44fed00/BB.jpg)

For this lab project, we are going to work with a list of blues musicians in various traditions.  You have the following files in your `lab06/data` directory:
    
    blues.txt
    swamppop.txt
    zydeco.txt

Open the file `blues.txt` in a text editor and examine its format.
    
    # Delta Blues
    Cecil Augusta
    Mose Allison
    Tommy Bankhead
    ...
    Elder Roma Wilson
    
    # Chicago Blues
    Alberta Adams
    Luther Allison
    ...

We are going to read the contents of this file and run it through the standard pipeline of data analysis:  data cleaning and preparation, data processing, and output.

### Data Cleaning and Preparation

A file simply contains an ordered collection of characters and digits.  Any logical interpretation is created by the user as he or she uses the data.  Thus we first need to ensure that incoming data are in a suitable format and structure for further analysis.

What does a suitable format look like?  If you were organizing a library of these musicians' works, you could sort by surname or by musical style.  There are several formats which could make sense, such as a collection of *database records* containing *fields* such as "artist name" and "musical style".  We will opt here for a spreadsheet-like organization:  three "columns" of data:  "surname", "first name", and "style".  Each entry will be a tuple inside of a list.

**Data Logical Format:**

    artist surname  artist first name   musical style
    Augusta         Cecil               Delta Blues
    Alexander       Alger "Texas"       Country Blues

**Data Machine Representation:**

    # musicians is a list of tuples
    musicians[0] = ('Augusta', 'Cecil', 'Delta Blues')
    musicians[1] = ('Alexander', 'Alger "Texas"', 'Country Blues')
    ...

Thus a single record should have one entry from each of these fields.  The data file we import from, however, is not in this format.  Thus when importing the data we have to:

1.  track which musical style we are currently importing; and
2.  *tokenize* the name into first-name and surname components; and
3.  add the tuple of these three items to a list.

Let's see how this works with a single record as a string.  If you have a line of text as a string,

```python
example = 'Jimmy Clanton'
print(example)
```

Output:
```
Jimmy Clanton
```

and you'd like to turn it into a tuple `record = (surname, first_name)`, how would you do it?

First, you'd probably want to split it into pieces so there is a separate `first_name` and `surname` to assign:

```python
names = example.split(' ')
print(names)
```

Output:
```
['Jimmy', 'Clanton']
```

Now you have a `list` called `names` which you can use to assign the separate variables in your tuple either in two steps:

```python
# step 1:  get the surname and first name separately
surname = names[1]
first_name = names[0]

# step 2:  populate the record tuple
record = (surname, first_name)
print(record)
```

Output:
```
('Clanton', 'Jimmy')
```

or equivalently in a single step:

```python
record = (names[1], names[0])
print(record)
```

Output:
```
('Clanton', 'Jimmy')
```

Since we're just using a simple tuple of strings, that's all you need to worry about to create a record.

This gets a little trickier with multiple first names, though:

```python
example = 'John Henry Barbee'
```

In this case, I suggest the following method.  First, split the name into its pieces like this:

```python
names = example.split(' ')
print(names)
```

Output:
```
['John', 'Henry', 'Barbee']
```

Next, pull the surname out as a single name.  (We have already removed Jrs. and Srs. from these data, so there's nothing exceptional left.)

```python
surname = names[-1]
print(surname)
```

Output:
```
Barbee
```

Finally, join the other names back together using the handy `join` method:

```python
first_name = ' '.join(names[:-1])  # note how this joins a list of strings by a blank space---it's an odd way of writing this
print(first_name)
```

Output:
```
John Henry
```

```python
record = (surname, first_name)
print(record)
```

Output:
```
('Barbee', 'John Henry')
```

<div class="alert alert-info">
Note that we are using tuples for the records.  Recall that tuples are like `list`s, but can't be changed (*immutable*).  Since we won't need to change a record once it's created, it makes sense in this case.
</div>

Let's now store a list of records.  Practice this process on the following entries:

```python
# Setup code for your exercise---do NOT edit this cell
records = [] # an empty list
entries = ['Ivory Joe Hunter', 'Etta James', 'Little Willie Littlefield', 'Robert Lowery', 'J. J. Malone', 'Percy Mayfield', 'Jimmy McCracklin']
```

<!-- grade_id: records, points: 0, grade: false, solution: true -->
```python
# Cycle through each musician name and add it to records in format.
for entry in entries:
    names = entry.split()# compose your code here
    surname = names[-1]
    first_name = ' '.join(names[:-1])# get the name elements out of entry
    # create a new record
    
    # append the record to the list records
    records.append( (surname,first_name) )
```

<!-- grade_id: records-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert records[0]  == ('Hunter', 'Ivory Joe')
assert records[-1] == ('McCracklin', 'Jimmy')
print('Success!')
```

Output:
```
Success!
```

Make sure you understand why there are two sets of parentheses on the last line (`records.append`), and ask a classmate or TA to explain if you do not.

Now we will turn the above process into a general script to load and process the data.

-   Compose a function `process` which accepts a string `filename`.  `process` should `return` a list of records contained in the file.

<!-- grade_id: process, points: 0, grade: false, solution: true -->
```python
# define your function here
def process(filename):
    entries = []# Create a blank list called `entries` and an empty string called `current_style`.
    current_style = ""
    myfile = open(filename)# Open the file `filename`, read the data, and close it.
    mydata = myfile.readlines()
    myfile.close()
    for line in mydata:
        data = line.strip()# Loop through each line of the file and do the following:
        if data == '':
            continue# Strip the whitespace off of the ends of the line using the `strip` method.
        
        # If the line is blank, `continue` execution.
        # (The `continue` statement makes Python just go back to the `for` loop again with the next value---
        # no more code is executed for the current value.)
        if data[0] == '#':
            data = data[1:]
            current_style = data
            continue
        else:
            name = data.strip().split()
            surname = name[-1]
            first_name = ' '.join(name[:-1])
            entries.append((surname,first_name,current_style))
            
        # If a line starts with `#`, it represents the musical style to be assigned until the next line with `#`.
        # In this case, remove the `'#'` from the beginning of the string and assign the musical style to `current_style`.
        # The loop should then `continue`.
        
        # Otherwise, a line contains a blues performer.  In this case, process the record much as you did above,
        # except that you also need to add a musical style to the tuple as the third element.
        
        # These data should be appended to the list `entries` as a single entry in the form of a tuple,
        #    (surname, first_name, current_style)
        
    return entries# Finally, `return` the list `entries`.
    pass # you can always delete a `pass` statement, since it does nothing
```

```python
filename = './data/blues.txt'
process(filename)# test your code here.  You may edit this cell, and you may use the files 'blues.txt', 'swamppop.txt', and 'zydeco.txt',
# all of which are located in the 'data/' directory.
```

Output:
```
[('Augusta', 'Cecil', 'Delta Blues'),
 ('Allison', 'Mose', 'Delta Blues'),
 ('Bankhead', 'Tommy', 'Delta Blues'),
 ('Barbee', 'John Henry', 'Delta Blues'),
 ('Bailey', 'Kid', 'Delta Blues'),
 ('Belfour', 'Robert', 'Delta Blues'),
 ('Booker', 'Charley', 'Delta Blues'),
 ('Bracey', 'Ishman', 'Delta Blues'),
 ('Brown', 'Willie', 'Delta Blues'),
 ('Burnside', 'R. L.', 'Delta Blues'),
 ('Carr', 'Sam', 'Delta Blues'),
 ('Carter', 'Bo', 'Delta Blues'),
 ('Cotton', 'James', 'Delta Blues'),
 ('Crudup', 'Arthur "Big Boy"', 'Delta Blues'),
 ('Billy', 'Delta Blind', 'Delta Blues'),
 ('Edwards', 'David Honeyboy', 'Delta Blues'),
 ('Hemphill', 'Jessie Mae', 'Delta Blues'),
 ('Hill', 'King Solomon', 'Delta Blues'),
 ('Hooker', 'John Lee', 'Delta Blues'),
 ('House', 'Son', 'Delta Blues'),
 ('Wolf', 'Howlin', 'Delta Blues'),
 ('James', 'Elmore', 'Delta Blues'),
 ('James', 'Skip', 'Delta Blues'),
 ('Johnson', 'Robert', 'Delta Blues'),
 ('Johnson', 'Tommy', 'Delta Blues'),
 ('Kimbrough', 'Junior', 'Delta
```

<!-- grade_id: process-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test basic single-genre case
zydeco = process('./data/zydeco.txt')
assert zydeco[0]  == ('Chavis', 'Boozoo', 'Zydeco')
assert zydeco[-3] == ('Dopsy', "Rockin'", 'Zydeco')
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: process-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test case with multiple genres
blues = process('./data/blues.txt')
assert blues[0]  == ('Augusta', 'Cecil', 'Delta Blues')
assert blues[-1] == ('Rose', 'Bayless', 'Piano Blues')
print('Success!')
```

Output:
```
Success!
```

Typically, you would also need to test for duplicates.  In this case, we've already removed duplicate entries that were present.  (This may occur because the many entries will have different musical styles associated with them—many musicians performed in several styles within the genre of blues.)

###  Data Processing

You should now have a function which loads a formated text file and converts that file into a collection of records which you can now use to ask and answer questions about the properties of the data set.

```python
blues = process('./data/blues.txt')
```

#### Sorting

For instance, sort by surname and list the first ten entries.

```python
blues.sort()  # note that sort sorts in place, rather than returning a result to you---this will trip you up if you are not careful!
blues[:10]
```

Output:
```
[('Abshire', 'Nathan', 'Swamp Blues'),
 ('Adams', 'Alberta', 'Chicago Blues'),
 ('Alexander', 'Alger "Texas"', 'Country Blues'),
 ('Alexander', 'Dave', 'West Coast Blues'),
 ('Alexander', 'Linsey', 'Chicago Blues'),
 ('Allison', 'Luther', 'Chicago Blues'),
 ('Allison', 'Mose', 'Delta Blues'),
 ('Ammons', 'Albert', 'Piano Blues'),
 ('Anderson', 'Pink', 'Piedmont Blues'),
 ('Armstrong', 'Howard "Louie Bluie"', 'Country Blues')]
```

Tuples will automatically be sorted by the first element, then the second, then the third; this makes sorting by surname easy.  Sorting by the other fields, such as first name or style, is a bit more involved.  You have to provide a `key` to `sort` so it knows what to sort by.  The easiest way is to write a function which `return`s the second element (or the first name, in this case):

```python
def second_element(a_list):
    return a_list[1]
blues = sorted(blues, key=second_element)
blues  # note that `"Baby Face" Leroy` sorts by `"` not `B`
```

Output:
```
[('Foster', '"Baby Face" Leroy', 'Chicago Blues'),
 ('Reed', 'A. C.', 'Chicago Blues'),
 ('Ammons', 'Albert', 'Piano Blues'),
 ('Collins', 'Albert', 'Chicago Blues'),
 ('Adams', 'Alberta', 'Chicago Blues'),
 ('Seward', 'Alec', 'Country Blues'),
 ('Alexander', 'Alger "Texas"', 'Country Blues'),
 ('Milburn', 'Amos', 'Piano Blues'),
 ('Odom', 'Andrew', 'Chicago Blues'),
 ('Stidham', 'Arbee', 'Chicago Blues'),
 ('Edwards', 'Archie', 'Piedmont Blues'),
 ('Burton', 'Aron', 'Chicago Blues'),
 ('Crudup', 'Arthur "Big Boy"', 'Delta Blues'),
 ('Spires', 'Arthur "Big Boy"', 'Chicago Blues'),
 ('Tate', 'Baby', 'Country Blues'),
 ('Bob', 'Barbecue', 'Country Blues'),
 ('Bob', 'Barbecue', 'Piedmont Blues'),
 ('Smith', "Barkin' Bill", 'Chicago Blues'),
 ('Chuck', 'Barrelhouse', 'Chicago Blues'),
 ('Rose', 'Bayless', 'Piano Blues'),
 ('Tucker', 'Bessie', 'Country Blues'),
 ('Broonzy', 'Big Bill', 'Chicago Blues'),
 ('Kinsey', 'Big Daddy', 'Chicago Blues'),
 ('McNeely', 'Big Jay', 'West Coast Blues'),

```

-   How many musical styles are there, and how many musicians in each style?
    
    To answer this question, we'll use a `dict`ionary.  A `dict` is like a `list`, except that a list is indexed by `int`s and a `dict` is indexed by many different data types.
    
    A `list` that contains strings, for example, can be indexed directly by the position of each string in the `list`:
    
    ![](./img/list-pic.png)

```python
list_example = ['alpha', 'beta', 'gamma', 'delta', 'omega']
print(list_example)
```

Output:
```
['alpha', 'beta', 'gamma', 'delta', 'omega']
```

A `dict`, in contrast, can use strings or floating-point values or tuples (and many other things) to index.  Consider a `dict` that uses English color names to identify HTML color codes:

![](./img/dict-pic.png)

```python
dict_example = {'red': '#FF0000', 'green': '#00FF00', 'yellow': '#FFFF00', 'blue': '#0000FF', 'black': '#000000'}
print(dict_example)
```

Output:
```
{'green': '#00FF00', 'yellow': '#FFFF00', 'red': '#FF0000', 'black': '#000000', 'blue': '#0000FF'}
```

```python
dict_example['yellow']
```

Output:
```
'#FFFF00'
```

You can add new entries to the `dict` very easily:

```python
dict_example['grey'] = '#888888'
```

So to count the number of musicians by style, we can use a `dict` to organize our counting.

1.  Loop over the records.
2.  For each record, get the musical style field.
3.  If the style is in the `dict`, add one to the count.
4.  If the style is *not* in the `dict`, add it (and give it a count of `1`).

This is how you see which keys (indices) are in the `dict`:

```python
list(dict_example.keys())
```

Output:
```
['green', 'grey', 'red', 'yellow', 'black', 'blue']
```

This is how you see if the `dict` has a particular key or not:

```python
'yellow' in dict_example.keys()
```

Output:
```
True
```

```python
'teal' in dict_example.keys()
```

Output:
```
False
```

-   Now that we've gotten this far with `dict`ionaries, let's revisit the original question:  how many artists are in each musical style?  Compose a function `count_styles` below which accepts a `list` of entries `records`.  `count_styles` should `return` a `dict` containing the musical styles represented and their respective count.

<!-- grade_id: count_styles, points: 0, grade: false, solution: true -->
```python
# define your function here
def count_styles(records):
    styles = {}  # a blank dictionary
    for record in records:
        sty = record[2]
        if sty not in styles:
            styles[sty] = 0
        if sty in styles:
            styles[sty] += 1
            
    return styles# Loop over the records.
        # For each record, get the musical style field.
        # If the style is in the dict, add one to the count.
        # i.e., styles['old style'] += 1
        
        # If the style is not in the dict, add it (and give it a count of 1).
        # i.e., styles['new style'] = 1
        
    # Return the resulting dict of styles
    pass # you can always delete a `pass` statement, since it does nothing
```

```python
zyd = process('./data/zydeco.txt')
zydecos = count_styles(zyd)
print(zydecos)# test your code here.  You may edit this cell, and you may use the files 'blues.txt', 'swamppop.txt', and 'zydeco.txt',
# all of which are located in the 'data/' directory.
```

Output:
```
{'Zydeco': 29}
```

<!-- grade_id: count_styles-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
zydeco = process('./data/zydeco.txt')
zydeco_styles = count_styles(zydeco)
assert zydeco_styles['Zydeco'] == 29
print('Success!')
zydeco_styles
```

Output:
```
Success!
{'Zydeco': 29}
```

<!-- grade_id: count_styles-test-blues, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
blues = process('./data/blues.txt')
blues_styles = count_styles(blues)
assert blues_styles['Gospel Blues'] == 19
print('Success!')
blues_styles
```

Output:
```
Success!
{'Chicago Blues': 148,
 'Country Blues': 51,
 'Delta Blues': 54,
 'Gospel Blues': 19,
 'Piano Blues': 31,
 'Piedmont Blues': 35,
 'Swamp Blues': 20,
 'West Coast Blues': 42}
```

So that tells you *how many times each thing occurs*.  What if you want to know *how many kinds of things there are*?

We need to make our list contain only unique elements—that is, remove multiple copies from it.  This is called *uniqifying* the list.

To uniqify your list, you can use the following code:

```python
# a short function to remove repeat elements of a list
def uniqify(input_list):
    # make an empty dictionary
    keys = {}
    for e in input_list:
        # add a key for each item in the list---duplicate keys will be overwritten
        keys[e] = 1
    return list(keys.keys())  # return a list of all unique keys
```

As an example, test this function on a simple list:

```python
my_list = [1,2,2,4,3,2,5,7,1,2]
uniqify(my_list)
```

Output:
```
[1, 2, 3, 4, 5, 7]
```

```python
surnames = []
for musician in blues:
   surnames.append(musician[0])

unique_surnames = uniqify(surnames)
unique_surnames.sort()
unique_surnames
```

Output:
```
['Abshire',
 'Adams',
 'Alexander',
 'Allison',
 'Ammons',
 'Anderson',
 'Armstrong',
 'Arnold',
 'Augusta',
 'B.',
 'Babe',
 'Bailey',
 'Baker',
 'Ball',
 'Band',
 'Bankhead',
 'Banks',
 'Barbee',
 'Barnes',
 'Bates',
 'Baty',
 'Becker',
 'Belfour',
 'Bell',
 'Belly',
 'Benoit',
 'Benton',
 'Billy',
 'Blackwell',
 'Blake',
 'Bloomfield',
 'Blue',
 'Bob',
 'Bogan',
 'Bonds',
 'Bonner',
 'Booker',
 'Boyd',
 'Bracey',
 'Bradley',
 'Branch',
 'Brim',
 'Brooks',
 'Broonzy',
 'Brown',
 'Brozman',
 'Buford',
 'Burnside',
 'Burton',
 'Butler',
 'Butterfield',
 'Campbell',
 'Carr',
 'Carroll',
 'Carter',
 'Caston',
 'Cephas',
 'Charles',
 'Charlie',
 'Chuck',
 'Clark',
 'Clarke',
 'Clayton',
 'Clearwater',
 'Coleman',
 'Collette',
 'Collins',
 'Cotten',
 'Cotton',
 'Council',
 'Cox',
 'Crayton',
 'Crudup',
 'Crutchfield',
 'Davenport',
 'Davis',
 'Dawkins',
 'DeSanto',
 'Diddley',
 'Dixon',
 'Dizz',
 'Domino',
 'Dorsey',
 'Doyle',
 'Drummer',
 'Duncan',
 'Dupree',
 'Edwards',
 'Estes',
 'Flynn
```

-   Now, use this ability to `uniqify` a list to determine how many different surnames there are.
    
    Compose a function `unique_surname_list` which accepts a `list` `records` containing the tuple entries.  `unique_surname_list` will `return` a `list` of unique surnames.

<!-- grade_id: surname_list, points: 0, grade: false, solution: true -->
```python
# define your function here
def unique_surname_list(records):
    surp = []
    for record in records:
        surp.append(record[0])
    unique_surp = uniqify(surp)
    unique_surp.sort()
    return unique_surp# your code here
    
    # Return the resulting list of surnames
    pass # you can always delete a `pass` statement, since it does nothing
```

```python
# test your code here.  You may edit this cell, and you may use the files 'blues.txt', 'swamppop.txt', and 'zydeco.txt',
# all of which are located in the 'data/' directory.
```

<!-- grade_id: surname_list-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
zydeco = process('./data/zydeco.txt')
zydeco_surnames = unique_surname_list(zydeco)
assert len(zydeco_surnames) == 24
print('Success!')
```

Output:
```
Success!
```

#### Tokenizing

To *tokenize* is to split a string into pieces (or *tokens*) by some rule.  For instance, you've done this with `split` before:

```python
"The Well at the World's End".split(' ')
```

Output:
```
['The', 'Well', 'at', 'the', "World's", 'End']
```

We can extract some statistical information about the naming of blues artists by tokenizing all components of their first names and then counting how many times each token (or bit) occurs.  For instance, `Blind Lemon Jefferson` became `('Jefferson', 'Blind Lemon', 'Gospel Blues')`; we now wish to tokenize `'Blind Lemon'` into `'Blind'` and `'Lemon'` (since `'Blind'` is a common moniker among Delta Blues performers and their musical descendants).

-   How many times does the name element `'Blind'` occur in the file `./data/blues.txt`?
    
    To find this, you will need to:
    
    1.  Tokenize all name bits.
    2.  Count the number of times each name bit occurs.  This is very similar to the counting of musical styles previously.

First, let's get all of the names (first and last) together in one `list`.

```python
blues = process('./data/blues.txt')
names = []
for musician in blues:
    names.append(musician[1])
    names.append(musician[0])
```

Next, tokenize each name and add the components to a master `list` of name bits.  Since you are adding a list of strings in each case to the list, you may wish to use the `extend` method instead of the `append` method.

```python
name_bits = []
for n in names:
    token = n.split()
    name_bits.append(token[0:2])# loop over names
# tokenize each name into name tokens
# DO NOT FORGET THE FIRST NAME AND THE SURNAME ENTRIES
# add the name tokens to names

name_bits
```

Output:
```
[['Cecil'],
 ['Augusta'],
 ['Mose'],
 ['Allison'],
 ['Tommy'],
 ['Bankhead'],
 ['John', 'Henry'],
 ['Barbee'],
 ['Kid'],
 ['Bailey'],
 ['Robert'],
 ['Belfour'],
 ['Charley'],
 ['Booker'],
 ['Ishman'],
 ['Bracey'],
 ['Willie'],
 ['Brown'],
 ['R.', 'L.'],
 ['Burnside'],
 ['Sam'],
 ['Carr'],
 ['Bo'],
 ['Carter'],
 ['James'],
 ['Cotton'],
 ['Arthur', '"Big'],
 ['Crudup'],
 ['Delta', 'Blind'],
 ['Billy'],
 ['David', 'Honeyboy'],
 ['Edwards'],
 ['Jessie', 'Mae'],
 ['Hemphill'],
 ['King', 'Solomon'],
 ['Hill'],
 ['John', 'Lee'],
 ['Hooker'],
 ['Son'],
 ['House'],
 ['Howlin'],
 ['Wolf'],
 ['Elmore'],
 ['James'],
 ['Skip'],
 ['James'],
 ['Robert'],
 ['Johnson'],
 ['Tommy'],
 ['Johnson'],
 ['Junior'],
 ['Kimbrough'],
 ['Little', 'Freddie'],
 ['King'],
 ['Robert'],
 ['Lockwood'],
 ['Willie'],
 ['Love'],
 ['Lead'],
 ['Belly'],
 ['Tommy'],
 ['McClennan'],
 ['Papa', 'Charlie'],
 ['McCoy'],
 ['Mississippi', 'Fred'],
 ['McDowell'],
 ['Mississippi', 'John'],
 ['Hurt'],
 ['Sonny', 'Boy'],
 ['Nelson'],
 
```

```python
name_bits = []
# loop over blues
for name in names:
    name_bits.extend(name.split(' '))
    
# tokenize each name into name tokens
# add the name tokens to names
name_bits
```

Finally, create a list of the count of each name bit as you did with `count_styles`.

```python
name_counts = {}  # a blank dictionary

for name_bit in name_bits:
    if name_bit in name_counts.keys():
        name_counts[name_bit] += 1
    else:
        name_counts[name_bit] = 1

name_counts
```

```python
# Sort the list and output it from most frequently occurring name bits to least.
def value(k):
    return name_counts[k]
name_count_list = sorted(name_counts, key=value)[::-1]

for name in name_count_list:
    print(name, name_counts[name])
```

It's now trivial for you to answer the original question of how many `'Blind'`s there are in the `blues.txt` file:

```python
print(name_counts['Blind'])
```

```python
# clear earlier data so we don't mask the behavior of the function `count_names` below
names = None
name_bits = None
name_count_list = None
name_counts = None
```

Go ahead and write this into a function `count_names`, which accepts a `list` `records` containing the tuple entries.  `count_names` will `return` a `dict` containing the number of times each name bit occurs.

<!-- grade_id: count_names, points: 0, grade: false, solution: true -->
```python
# define your function here
def count_names(records):
    name_counts = {}  # a blank dictionary
    
    # loop over names in records
        # tokenize each name into name tokens---DO NOT FORGET THE FIRST NAME AND THE SURNAME ENTRIES
        # add the name tokens to names
    
    # create a list of the count of each name bit as you did with count_styles.
    
    # sort the list from most frequently occurring name bits to least
    
    # return the resulting dict counting name bits.
    pass # you can always delete a `pass` statement, since it does nothing
```

```python
# test your code here.  You may edit this cell, and you may use the files 'blues.txt', 'swamppop.txt', and 'zydeco.txt',
# all of which are located in the 'data/' directory.
```

<!-- grade_id: count_names-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test number of name elements found
blues = process('./data/blues.txt')
blues_names = count_names(blues)
assert len(list(blues_names.keys())) == 594
print('Success!')
```

<!-- grade_id: count_names-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test success in counting name elements
blues = process('./data/blues.txt')
blues_names = count_names(blues)
assert blues_names['Wheatstraw'] == 1
assert blues_names['Johnny'] == 14
print('Success!')
```

# Before you submit...

Before you submit this exercise, make sure that everything runs as expected.  You can either use the *Validate* button on the *Assignments* tab or type `nbgrader validate labX` at the command line.  **You should also *save* before submitting.**

Make sure that you have filled in any code block which says "YOUR CODE HERE" or "YOUR ANSWER HERE".

List any resources used (beyond the TA or the help files):

This laboratory exercise was written by Neal Davis for Computer Science 101 at the University of Illinois.

©2015–16 University of Illinois

![](./cs101-footer.png)